```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

# Lab 8 – Feature Engineering ⚙️

## DSC 80, Fall 2023

### Due Date: Monday, November 27th at 11:59pm

## Instructions
Welcome to the eighth lab assignment in DSC 80 this quarter!

Much like in DSC 10, this Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. Unlike DSC 10, the notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:
- **Unlike in DSC 10, labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the [course GitHub repository](https://github.com/dsc-courses/dsc80-2023-fa).
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- **To ensure that all of your work to be submitted is in `lab.py`, we've provided an additional uneditable notebook, called `lab-validation.ipynb`, that contains only the tests and their setup. Make sure you are able to run it top-to-bottom without error before submitting!**
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**Importing code from `lab.py`**:

* Below, we import the `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.
    
***Note***: In addition to the usual imports below, we'll need the `statsmodels` package that isn't included in the standard `dsc80` conda environment we had you configure. (We need it for some of the `plotly` graphs created in Part 1.)

Before running the cell below for the first time, run the following line in your Terminal after activating your `dsc80` conda environment, and then restart your kernel. If you've already installed `statsmodels`, you should be able to run this cell without error.

```
    pip install statsmodels
```

```python
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import plotly.express as px
import statsmodels.api as sm
from pathlib import Path
from sklearn.preprocessing import Binarizer, QuantileTransformer, FunctionTransformer
import itertools
import warnings
warnings.filterwarnings('ignore')
```

## Part 1: Scaling Transformations 📐

A scaling transformation transforms the scale of the data of a particular quantitative column. Mathematically, each data point $d_i$ is replaced with a transformed value $t_i = f(d_i)$, where $f$ is a transformation function. We can transform any column in a dataset, whether it corresponds to a feature or a target.

Generally, the goal of a scaling transformation is to change the data from a complicated, non-linear relationship into a **linear** relationship. Linear relationships are very easy to understand and are easily used by models, like linear regression.

However, non-linear growth is a commonly seen relationship in data. Sometimes this growth is by a **fixed power** and sometimes it is **exponential**. The scaling transformations that turn these types of growth linear are **root** and **log** transformations respectively. (Generally, it is more difficult to determine which transformation is appropriate for a given dataset, though the [Tukey-Mosteller bulge diagram](https://freakonometrics.hypotheses.org/files/2014/06/Selection_005.png) is useful.)

Let's start by looking at some examples of transformations.

#### Example 1

Run the cell below to generate a scatter plot.

```python
# By setting a seed, we guarantee that we will see the same results each time we run this cell.
np.random.seed(23)

# Generates a random scatter plot
x = np.arange(1, 101) + np.random.normal(0, 0.5, 100)
y = 2 * ((x + np.random.normal(0, 1, 100)) ** 2) + np.abs(x) * np.random.normal(0, 30, 100)
df_1 = pd.DataFrame().assign(x=x, y=y)

px.scatter(df_1, x='x', y='y', trendline="ols", trendline_color_override="red")
```

It doesn't appear to be the case that `'x'` and `'y'` are linearly associated here, and they aren't – there is a **quadratic** relationship between them. Note that if we were to create a **residual plot** above, there would be a pattern – the residuals for smaller `'x'` would mostly be positive, and the residuals for larger `'x'` would mostly be negative. From [DSC 10](https://inferentialthinking.com/chapters/15/5/Visual_Diagnostics.html), we know that patterns in a residual plot imply that the relationship between the two variables is non-linear.

To linearize the relationship, we can take the square root of each `'y'` value:

```python
df_1['root y'] = np.sqrt(df_1['y'])

px.scatter(df_1, x='x', y='root y', trendline="ols", trendline_color_override="red")
```

That looks much better!

#### Example 2

Run the cell below to generate another scatter plot.

```python
# By setting a seed, we guarantee that we will see the same results each time we run this cell
np.random.seed(32)

# Generates a different random scatter plot
x = np.linspace(2, 5, 100)
y = 10 * (np.e ** x) + np.abs(x) * np.random.normal(0, 5, 100) + np.random.normal(0, 30, 100)
df_2 = pd.DataFrame().assign(x=x, y=y)

px.scatter(df_2, x='x', y='y', trendline="ols", trendline_color_override="red")
```

Again, the relationship between `'x'` and `'y'` is not quite linear. Let's try the square root transformation we tried in Example 1:

```python
df_2['root y'] = np.sqrt(df_2['y'])

px.scatter(df_2, x='x', y='root y', trendline="ols", trendline_color_override="red")
```

Hmm... the relationship certainly looks _more_ linear than before, but still not quite linear. Let's look at the residual plot:

```python
# Feel free to use this function directly to help you answer Question 1.
def create_residual_plot(df, x, y):
    df = df.copy()
    from sklearn.linear_model import LinearRegression
    model = LinearRegression()
    model.fit(df[[x]], df[y])
    df['pred'] = model.predict(df[[x]])
    df[f'{y} residuals'] = df[y] - model.predict(df[[x]])
    return px.scatter(df, x='pred', y=f'{y} residuals', trendline='ols', trendline_color_override='red')

create_residual_plot(df_2, 'x', 'root y')
```

There is clearly a pattern in the residual plot. Let's instead try another transformation for the `'y'` values – $\log$.

```python
df_2['log y'] = np.log(df_2['y'])

px.scatter(df_2, x='x', y='log y', trendline="ols", trendline_color_override="red")
```

That looks much better! We can verify that the residual plot has no "patterns":

```python
create_residual_plot(df_2, 'x', 'log y')
```

Note – there is still evidence of **heteroscedasticity**, or "uneven spread", in this scatter plot, but the relationship is as close to linear as we'll get.

### Question 1 –  Root vs. Log

Now that we've learned how to perform transformations with example datasets, it's your job to apply these ideas to a real dataset. Below, you are given a dataset that describes the [number of home runs in the MLB per year](https://www.mlb.com/glossary/standard-stats/home-run). The relationship between the two variables, `'Year'` and `'Homeruns'`, is not linear.

**Specifically, your job is to determine what the appropriate transformation to apply to the `'Home runs'` column is, in order to linearize the relationship.**

Complete the implementation of the function `best_transformation`, which returns either 1, 2, 3, or 4, with the value corresponding to one of the following choices:

1. Square root transformation.
2. Log transformation.
3. Both work the same.
4. Neither gives a transformation revealing a linear relationship.

***Hint***: If you find that both residual plots have some sort of pattern, choose the residual plot in which the vertical spread is constant. There is one clearly correct answer.

```python
homeruns_fp = Path('data')/'homeruns.csv'
homeruns = pd.read_csv(homeruns_fp)
```

```python
x = homeruns['Year']
y = homeruns['Homeruns']
homeruns_copy = pd.DataFrame().assign(x=x, y=y)

px.scatter(homeruns_copy, x='x', y='y', trendline="ols", trendline_color_override="red")
```

```python
homeruns_copy['root y'] = np.sqrt(homeruns_copy['y'])

px.scatter(df_2, x='x', y='root y', trendline="ols", trendline_color_override="red")
```

```python
create_residual_plot(homeruns_copy, 'x', 'root y')
```

```python
homeruns_copy['log y'] = np.log(homeruns_copy['y'])

px.scatter(homeruns_copy, x='x', y='log y', trendline="ols", trendline_color_override="red")
```

```python
create_residual_plot(homeruns_copy, 'x', 'log y')
```

```python
grader.check("q1")
```

## Part 2: Diamond Pricing 💎

In this next section, you will pretend you are a jewelry appraiser and predict the prices of diamonds given several standard characteristics of diamonds.

You will use linear regression to predict prices, while improving the quality of your predictions using **feature engineering**. Since this question is supposed to help you understand feature engineering, **you will be building these features from scratch, instead of using the built in `sklearn` or `pandas` methods**.

The `diamonds` dataset is accessible via `seaborn` (with `sns.load_dataset('diamonds')`), but we've skipped that step and loaded it for you below. The DataFrame has 53940 rows and 10 columns:

|column|description|unique values or range|
|---|---|---|
|`'carat'`|weight of the diamond in carats (each carat is 0.2 grams)| 0.2 - 5.01 |
|`'cut'`|quality of the cut | Fair, Good, Very Good, Premium, Ideal |
|`'color'`|diamond colour | J (worst, near colorless), I, H, G, F, E, D (best, absolute colorless) |
|`'clarity'`|a measurement of how clear the diamond is | I1 (worst), SI2, SI1, VS2, VS1, VVS2, VVS1, IF (best) |
|`'depth'`|total depth percentage, computed as z / mean(x, y) = 2 * z / (x + y) | 43 - 79 |
|`'table'`|width of top of diamond relative to widest point | 43 - 95 |
|`'price'`|price in US dollars | \\$326 - \\$18,823 USD |
|`'x'`|length in mm | 0 - 10.74 |
|`'y'`|width in mm | 0 - 58.9 | 
|`'z'`|depth in mm | 0 - 31.8 |

If you want to learn more about how diamonds are measured, refer to [this page by the American Gem Society](https://www.americangemsociety.org/4cs-of-diamonds/).

```python
diamonds = pd.read_csv(Path('data')/'diamonds.csv')
diamonds.head()
```

Output:
```
carat      cut color clarity  depth  table  price     x     y     z
0   0.23    Ideal     E     SI2   61.5   55.0    326  3.95  3.98  2.43
1   0.21  Premium     E     SI1   59.8   61.0    326  3.89  3.84  2.31
2   0.23     Good     E     VS1   56.9   65.0    327  4.05  4.07  2.31
3   0.29  Premium     I     VS2   62.4   58.0    334  4.20  4.23  2.63
4   0.31     Good     J     SI2   63.3   58.0    335  4.34  4.35  2.75
```

### Question 2 – Ordinal Encoding 🔢

Every categorical variable in the dataset is an ordinal column, meaning that there is an inherent order that we can use to sort the values in the column. Recall that **ordinal encoding** is a feature transformation that maps the values in an ordinal column to positive integers in a way that preserves the order of the column values. For instance, an ordinal encoding for Freshman, Sophomore, Junior, Senior is 0, 1, 2, 3.

Complete the implementation of the function `create_ordinal`, which takes in the `diamonds` DataFrame and returns a DataFrame of ordinal features only with names of the form `'ordinal_<col>'`, where `'<col>'` is the original categorical column name. For instance, the `'ordinal_color'` column should consist of values from 0 to 6, where 0 refers to `'J'` and 6 refers to `'D'`. (In all cases, start counting from 0.)

***Notes:*** 
- Remember, you are creating this function using basic `pandas`. You might want to create a helper function that takes in a single column and an ordering for that column.
- Don't include non-ordinal features in the returned DataFrame. That is, if there are only three columns in `diamonds` that are ordinal, `create_ordinal` should return a DataFrame with three columns.
- The orderings for each of the ordinal columns are displayed in the data dictionary above (in the `'unique values or range'` column).

```python
def quality_cut(str_in):
  if str_in == 'Fair':
    return 0
  elif str_in == 'Good':
    return 1
  elif str_in == 'Very Good':
    return 2
  elif str_in == 'Premium':
    return 3
  else:
    return 4
```

```python
def color(str_in):
  if str_in == 'J':
    return 0
  elif str_in == 'I':
    return 1
  elif str_in == 'H':
    return 2
  elif str_in == 'G':
    return 3
  elif str_in == 'F':
    return 4
  elif str_in == 'E':
    return 5
  else:
    return 6
```

```python
def clarity(str_in):
  if str_in == 'I1':
    return 0
  elif str_in == 'SI2':
    return 1
  elif str_in == 'SI1':
    return 2
  elif str_in == 'VS2':
    return 3
  elif str_in == 'VS1':
    return 4
  elif str_in == 'VVS2':
    return 5
  elif str_in == 'VVS1':
    return 6
  else:
    return 7
```

```python
diamonds_copy = diamonds[['cut', 'color', 'clarity']]
diamonds_copy['cut'] = diamonds_copy['cut'].apply(quality_cut)
diamonds_copy['color'] = diamonds_copy['color'].apply(color)
diamonds_copy['clarity'] = diamonds_copy['clarity'].apply(clarity)
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(Path('data')/'diamonds.csv')
out_q2 = create_ordinal(diamonds)
```

```python
set(out_q2.columns)
```

Output:
```
{'ordinal_clarity', 'ordinal_color', 'ordinal_cut'}
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### Question 3 – Nominal Encoding 📊

Even though the categorical variables in the dataset are ordinal, we can still treat them as nominal by forgetting their order. To treat the categorical variables in our dataset as nominal, we might **one-hot encode** them. 

#### `create_one_hot`

Complete the implementation of the function `create_one_hot`, which takes in the `diamonds` DataFrame and returns a DataFrame of one-hot encoded features with names of the form `'one_hot_<col>_<val>'`, where `'<col>'` is the original categorical column name, and `'<val>'` is the value found in the categorical column `'<col>'`. For instance, one of your column names will be `'one_hot_color_J'`.

***Notes:***
- Only include one-hot-encoded columns in the DataFrame that `create_one_hot` returns.
- Create a helper function that creates the one-hot encoding for a single column. **Do not** use `sklearn` or `pd.get_dummies` for this question!
- As per usual, write an efficient implementation. You may use a `for`-loop over **columns**, but not over rows. And the order of **columns** does not matter.
- In lecture, we will look at cases where we need to drop one one-hot encoded column per categorical variable. **Do not drop** any one-hot encoded columns here!

<br>

#### `create_proportions`

Similar to the one-hot encoding case, you can replace a value in a nominal column with the proportion of times that value appears in the column. For instance, if a column consists of the values `['a', 'b', 'a', 'c']`, then the proportion-encoded column is `[0.5, 0.25, 0.5, 0.25]`.  This might be a reasonable approach to predicting the price of a diamond, as you might expect *rarer attributes to be considered more valuable* than common ones.

Complete the implementation of the function `create_proportions`, which takes in the `diamonds` DataFrame and returns a DataFrame of proportion-encoded features with names of the form `'proportion_<col>'`, where `'<col>'` is the original categorical column name.

```python
diamonds
```

Output:
```
carat        cut color clarity  depth  table  price     x     y     z
0       0.23      Ideal     E     SI2   61.5   55.0    326  3.95  3.98  2.43
1       0.21    Premium     E     SI1   59.8   61.0    326  3.89  3.84  2.31
2       0.23       Good     E     VS1   56.9   65.0    327  4.05  4.07  2.31
3       0.29    Premium     I     VS2   62.4   58.0    334  4.20  4.23  2.63
4       0.31       Good     J     SI2   63.3   58.0    335  4.34  4.35  2.75
...      ...        ...   ...     ...    ...    ...    ...   ...   ...   ...
53935   0.72      Ideal     D     SI1   60.8   57.0   2757  5.75  5.76  3.50
53936   0.72       Good     D     SI1   63.1   55.0   2757  5.69  5.75  3.61
53937   0.70  Very Good     D     SI1   62.8   60.0   2757  5.66  5.68  3.56
53938   0.86    Premium     H     SI2   61.0   58.0   2757  6.15  6.12  3.74
53939   0.75      Ideal     D     SI2   62.2   55.0   2757  5.83  5.87  3.64

[53940 rows x 10 columns]
```

```python
def nominal_col(ordin, df, col):
    for val in df[col].unique():
        ordin[f'one_hot_{col}_{val}'] = (df[col] == val).astype(int)
    return ordin
```

```python
dfin = pd.DataFrame()
dfin
```

Output:
```
Empty DataFrame
Columns: []
Index: []
```

```python
nominal_col(dfin, diamonds, 'cut')
```

Output:
```
one_hot_cut_Ideal  one_hot_cut_Premium  one_hot_cut_Good  \
0                      1                    0                 0   
1                      0                    1                 0   
2                      0                    0                 1   
3                      0                    1                 0   
4                      0                    0                 1   
...                  ...                  ...               ...   
53935                  1                    0                 0   
53936                  0                    0                 1   
53937                  0                    0                 0   
53938                  0                    1                 0   
53939                  1                    0                 0   

       one_hot_cut_Very Good  one_hot_cut_Fair  
0                          0                 0  
1                          0                 0  
2                          0                 0  
3     
```

```python
def create_one_hot(df):
    ordin_df = pd.DataFrame()
    for col in ['cut', 'color', 'clarity']:
        ordin_df = nominal_col(ordin_df, df, col)
    return ordin_df
```

```python
create_one_hot(diamonds)
```

Output:
```
one_hot_cut_Ideal  one_hot_cut_Premium  one_hot_cut_Good  \
0                      1                    0                 0   
1                      0                    1                 0   
2                      0                    0                 1   
3                      0                    1                 0   
4                      0                    0                 1   
...                  ...                  ...               ...   
53935                  1                    0                 0   
53936                  0                    0                 1   
53937                  0                    0                 0   
53938                  0                    1                 0   
53939                  1                    0                 0   

       one_hot_cut_Very Good  one_hot_cut_Fair  one_hot_color_E  \
0                          0                 0                1   
1                          0                 0                1   
2
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(Path('data')/'diamonds.csv')
out1_q3 = create_one_hot(diamonds)
out2_q3 = create_proportions(diamonds)
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

### Question 4 – Quadratic Features 📈

Linear regression doesn't capture non-linear relationships between variables. However, you can create features that encode such dependencies **before** fitting your regression model. Creating polynomial features is one way to do this. For example, the diamonds dataset contains `'x'`, `'y'`, and `'z'` dimensions for each stone. However, different combinations of size may be more valuable than others: a "deep and wide" diamond might be considered more valuable than a shallow, but "long and wide" diamond.

Complete the implementation of the function `create_quadratics`, which takes in the `diamonds` DataFrame and returns a DataFrame of quadratic features of the form `'<col1> * <col2>'`, where `'<col1>'` and `'<col2>'` are the original quantitative columns. The output DataFrame should contain a column for every distinct pair of quantitative columns in `diamonds` (aside from `price`, which should be left out as it is what we are predicting). For instance, one of the columns in the returned DataFrame should named either `'carat * x'` or `'x * carat'`; the order of column names is not important.

***Notes:***
- Again, **do not** use `sklearn` for this question! 
- Try finding all pairs of quantitative columns efficiently; don't use a nested loop (hint: you may `import itertools`). Our solution contains just a single `for`-loop (over pairs of columns).
- If you import itertools, you must include the itertools import in both this notebook and in `lab.py`.
- The columns of the resulting DataFrame may be in any order.

```python
diamonds
```

Output:
```
carat        cut color clarity  depth  table  price     x     y     z
0       0.23      Ideal     E     SI2   61.5   55.0    326  3.95  3.98  2.43
1       0.21    Premium     E     SI1   59.8   61.0    326  3.89  3.84  2.31
2       0.23       Good     E     VS1   56.9   65.0    327  4.05  4.07  2.31
3       0.29    Premium     I     VS2   62.4   58.0    334  4.20  4.23  2.63
4       0.31       Good     J     SI2   63.3   58.0    335  4.34  4.35  2.75
...      ...        ...   ...     ...    ...    ...    ...   ...   ...   ...
53935   0.72      Ideal     D     SI1   60.8   57.0   2757  5.75  5.76  3.50
53936   0.72       Good     D     SI1   63.1   55.0   2757  5.69  5.75  3.61
53937   0.70  Very Good     D     SI1   62.8   60.0   2757  5.66  5.68  3.56
53938   0.86    Premium     H     SI2   61.0   58.0   2757  6.15  6.12  3.74
53939   0.75      Ideal     D     SI2   62.2   55.0   2757  5.83  5.87  3.64

[53940 rows x 10 columns]
```

```python
def create_quadratics(df):
    # Exclude non-quantitative columns or columns not to be used for creating quadratics.
    # Assuming 'price' is not included in the given DataFrame 'df'.
    quantitative_columns = df.select_dtypes(include=[float, int]).columns.tolist()

    # Create a list of all possible pairs for quadratic features
    # excluding pairs that include the 'price' column.
    pairs = list(itertools.combinations(quantitative_columns, 2))

    # Initialize an empty DataFrame to store the quadratic features
    quadratics = pd.DataFrame(index=df.index)

    # Iterate through each pair and calculate the product (quadratic feature)
    for (col1, col2) in pairs:
        quadratics[f'{col1} * {col2}'] = df[col1] * df[col2]

    return quadratics

# Apply the function to the mock-up DataFrame
quadratic_features = create_quadratics(diamonds)
quadratic_features
```

```python
quantitative_columns = diamonds.select_dtypes(include=[float, int]).columns.tolist()
quantitative_columns.remove('price')
pairs = list(itertools.combinations(quantitative_columns, 2))
result = pd.DataFrame()
for (col1, col2) in pairs:
    result[f'{col1} * {col2}'] = diamonds[col1] * diamonds[col2]
result
```

Output:
```
carat * depth  carat * table  carat * x  carat * y  carat * z  \
0             14.145          12.65     0.9085     0.9154     0.5589   
1             12.558          12.81     0.8169     0.8064     0.4851   
2             13.087          14.95     0.9315     0.9361     0.5313   
3             18.096          16.82     1.2180     1.2267     0.7627   
4             19.623          17.98     1.3454     1.3485     0.8525   
...              ...            ...        ...        ...        ...   
53935         43.776          41.04     4.1400     4.1472     2.5200   
53936         45.432          39.60     4.0968     4.1400     2.5992   
53937         43.960          42.00     3.9620     3.9760     2.4920   
53938         52.460          49.88     5.2890     5.2632     3.2164   
53939         46.650          41.25     4.3725     4.4025     2.7300   

       depth * table  depth * x  depth * y  depth * z  table * x  table * y  \
0             3382.5    242.925    244.770    149.445     217.2
```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(Path('data')/'diamonds.csv')
out_q4 = create_quadratics(diamonds)
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

### Question 5 – Comparing Performance 🏆

We've now created several sets of features. **Which features are best able to predict the price of a diamond in a linear regression model?** We'll look at both single-feature linear regression models and multiple regression models. In all cases, use the default arguments to `sklearn`'s `LinearRegression` object (i.e. assume there is an intercept term).

<br>

Two of the evaluation metrics we can use to compare regression models are RMSE and $R^2$.

RMSE, or Root Mean Squared Error, roughly measures the average distance between a model's predicted values and the actual values. The RMSE indicates how concentrated predictions are around the line of best fit, with lower values typically indicating smaller residuals and better model performance. RMSE can be calculated using the formula below, where $\text{n}$ is the number of data points, $\hat{y_i}$ is the $i$th prediction, and $y_i$ is the corresponding actual value.

$$\text{RMSE} = \sqrt{\frac{\sum_{i=1}^n (y_i - \hat{y_i})^2}{n}}$$

Remember that in least squares regression, we choose the model parameters that minimize RMSE.

<br>

$R^2$, also known as the coefficient of determination, denotes the proportion of variance in the dependent variable that the independent variable explains. It is used to assess the quality of a linear fit, and in the simple linear regression case, it is the square of the correlation coefficient $r$. $R^2$ values range between 0 and 1, with lower values indicating a worse linear fit and with higher values indicating a better linear fit. If `model` is a fitted `sklearn` regression object, $R^2$ can be computed by calling `model.score`. 

#### Single-Feature Models

- (1) Fit a single-feature linear regression model on `'carat'`. What is the $R^2$ of the model? (Note that `'carat'` turns out to be the best single feature to use in a linear model that predicts price.)
- (2) What is the RMSE of the model you created in (1)?
- (3) Amongst the other **quantitative** features present in the original `diamonds`, which produces the single-feature linear regression model with the highest $R^2$?
- (4) Amongst all the new features you created in Questions 2-4, which produces the single-feature linear regression model with the highest $R^2$?
- (5) Amongst the new categorical features you created in Question 2 and 3, which produces the single-feature linear regression model with the highest $R^2$? 

#### Multiple Regression

Now, fit a multiple regression model using:
- the quantitative columns that were present in the original `diamonds` dataset, and
- the quantitative features engineered in Question 4

as features. (Don't use any of the encodings of categorical columns from Questions 2 and 3.)

- (6) What is the RMSE of this new model?


<br>

Complete the implementation of the function `comparing_performance`, which returns a list containing the answers to the 6 questions numbered (1), (2), ..., (6) above. You don't need to round any of your answers.

***Hint:*** Repeatedly use the `sklearn` pattern included below. It's a good idea to make a helper function that takes in a column, performs single-feature regression using the input column as the feature, and returns the $R^2$ and RMSE of the model.

```python
from sklearn.linear_model import LinearRegression

# X = ...
# y = ...

# lr = LinearRegression()
# lr.fit(X, y)  # X is a DataFrame of training data; y is a Series of prices
# lr.score(X, y)  # R-squared
# lr.predict(X) # predicted prices
```

```python
def Rsquare_RMSE(x, y):
    X = x
    y = y
    lr = LinearRegression()
    lr.fit(X, y)  # X is a DataFrame of training data; y is a Series of prices
    R = lr.score(X, y)  # R-squared
    predicted = lr.predict(X) # predicted prices
    RMSE = np.sqrt(sum((y - predicted)**2) / x.shape[0])
    return (R, RMSE)
```

```python
models = []
for col in ['carat', 'depth', 'table', 'x', 'y', 'z']:
    RS, RMSE = Rsquare_RMSE(diamonds[[col]], diamonds['price'])
    models.append([col, RS, RMSE])
models
```

Output:
```
[['carat', 0.8493305264354858, 1548.5331930613002],
 ['depth', 0.00011336722437849112, 3989.1766174607033],
 ['table', 0.016163029068700818, 3957.0310019807202],
 ['x', 0.782225554041623, 1861.7070454430955],
 ['y', 0.7489533304600555, 1998.872603829358],
 ['z', 0.7417506045344293, 2027.3444398443723]]
```

```python
df = pd.DataFrame(models, columns=['vars', 'RS', 'RMSE']).set_index('vars')
df
output = [df['RS'].loc['carat'], df['RMSE'].loc['carat'], df['RS'].drop(['carat']).idxmax()]
output
```

Output:
```
[0.8493305264354858, 1548.5331930613002, 'x']
```

```python
out_q2 = create_ordinal(diamonds)
for col in out_q2.columns:
    RS, RMSE = Rsquare_RMSE(out_q2[[col]], diamonds['price'])
    models.append([col, RS, RMSE])
out_q3_hot = create_one_hot(diamonds)
for col in out_q3_hot.columns:
    RS, RMSE = Rsquare_RMSE(out_q3_hot[[col]], diamonds['price'])
    models.append([col, RS, RMSE])
out_q3_prop = create_proportions(diamonds)
for col in out_q3_prop.columns:
    RS, RMSE = Rsquare_RMSE(out_q3_prop[[col]], diamonds['price'])
    models.append([col, RS, RMSE])
df3 = pd.DataFrame(models, columns=['vars', 'RS', 'RMSE']).set_index('vars')
```

```python
df3['RS'].drop(['carat', 'depth', 'table', 'x', 'y', 'z']).idxmax()
```

Output:
```
'ordinal_color'
```

```python
df2 = pd.DataFrame(models, columns=['vars', 'RS', 'RMSE']).set_index('vars')
df2
output.append(df['RS'].idxmax())
out_q4 = create_quadratics(diamonds)
for col in out_q4.columns:
    RS, RMSE = Rsquare_RMSE(out_q4[[col]], diamonds['price'])
    models.append([col, RS, RMSE])
df = pd.DataFrame(models, columns=['vars', 'RS', 'RMSE']).set_index('vars')
output.append(df['RS'].idxmax())
output
```

Output:
```
[0.8493305264354858,
 1548.5331930613002,
 'x',
 'carat',
 'carat * x',
 'carat * x',
 'carat * x']
```

```python
#Multi
X = pd.concat([diamonds[['carat', 'depth', 'table', 'x', 'y', 'z']], out_q4], axis=1)
output.append(Rsquare_RMSE(X, diamonds['price'])[1])
output
```

Output:
```
[0.8493305264354858,
 1548.5331930613002,
 'x',
 1434.840008904737,
 1434.840008904737]
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
import numbers
out_q5 = comparing_performance()
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

## Part 3 – Feature Engineering with `sklearn` 🧠

In this final question, you will use `sklearn`'s transformers and estimators for feature engineering. While everything you do with `sklearn` is possible to do with `pandas`, `sklearn` transformers enable you to couple your feature engineering with your modeling. This will allow you to more quickly build and assess your models in `sklearn`.

Specifically, you will create a `TransformDiamonds` class that contains the three methods specified below – `transform_carat` (6.1), `transform_to_quantile` (6.2), and `transform_to_depth_pct` (6.3). In the starter code, there is a skeleton for `TransformDiamonds` that is initialized with a DataFrame `diamonds`.

Each of the methods you implement in the `TransformDiamonds` class should take in a DataFrame, initialize a specific `sklearn.Transformer` object (like `Binarizer` or `FunctionTransformer`), and use the transformer to transform columns from the input DataFrame. You should **not** use DataFrame methods like `apply` in this problem.

```python
from sklearn.preprocessing import Binarizer, QuantileTransformer, FunctionTransformer
```

Question 6 is made up of the three subparts below.

### Question 6.1 – Transforming a Quantitative Column into a Binary Column (`transform_carat`)

We call a diamond **large** if its weight is strictly greater than 1 carat. We want to **binarize** weights, so that they are 1 for large diamonds and 0 for small diamonds. Complete the implementation fo the method `transform_carat`, which takes in a DataFrame like `diamonds` and returns a binarized **array** of weights. Use a `Binarizer` object as your transformer.

***Notes:***
- You will return an array, not a Series, because `sklearn` thinks in terms of `np.ndarray`s, not DataFrames.
- The implementation of this function should only take two lines.

```python

```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(Path('data')/'diamonds.csv')
q6a_trans = TransformDiamonds(diamonds)
q6a_out = q6a_trans.transform_carat(diamonds)
```

```python
grader.check("q6.1")
```

Output:
```
q6.1 results: All test cases passed!
```

### Question 6.2 – Transforming a Quantitative Column into Quantiles (`transform_to_quantile`)

You will now transform the `'carat'` column so that each diamond's weight in carats is replaced with the **percentile** amongst all diamonds in which its weight lies.

Complete the implementation of the method `transform_to_quantiles`, which takes in a DataFrame like `diamonds` and returns an array containing the percentiles of the weight of each diamond, amongst all diamonds in `self.data`. This array should consist of proportions between 0 and 1; for instance, 0.65 will refer to the 65th percentile. The relevant transformer is `QuantileTransformer`. 

Some guidance:

- Unlike with `Binarizer`, you need to `fit` your `QuantileTransformer` before calling `transform` on the input DataFrame `data`. 
    - You should `fit` your transformer on the DataFrame `self.data`, but you should only `transform` the `data` that is passed to `transform_to_quantiles`. 
    - Note that these two DataFrames, `self.data` and `data`, don't have to be the same! For instance, in the last two lines of the testing setup cell below, we fit a `QuantileTransformer` using just the first 1000 rows of `diamonds`, and then `transform` the entire `diamonds` DataFrame. Make sure your `transform_to_quantiles` method works in such a case.
- When initializing your `QuantileTransformer`, use `n_quantiles=100`.

```python

```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
q6b_trans = TransformDiamonds(diamonds)
q6b_out = q6b_trans.transform_to_quantile(diamonds)
q6b_trans_top_1000 = TransformDiamonds(diamonds[:1000])
q6b_out_top_1000 = q6b_trans_top_1000.transform_to_quantile(diamonds)
```

```python
grader.check("q6.2")
```

Output:
```
q6.2 results: All test cases passed!
```

### Question 6.3 – Transforming a Quantitative Column Using a Formula

Recall from the introduction to Part 2 that the "depth percentage" of a diamond is defined as

$$\text{Depth Pct.} = 100\% \cdot \frac{2z}{x + y}$$

where $x$, $y$, and $z$ come from the `'x'`, `'y'`, and `'z'` columns in `diamonds`.

Let's suppose that for some reason we don't have access to the `'depth'` column in `diamonds`, and instead need to recreate it just by looking at the `'x'`, `'y'`, and `'z'` columns. 

Complete the implementation of the method `transform_to_depth_pct`, which takes in a DataFrame like `diamonds` and returns an array consisting of the depth percentages of each diamond. Percentages should be between 0 and 100. The relevant transformer is `FunctionTransformer`.

***Notes:***
- To use `FunctionTransformer`, you will need to define your own function that takes in a 2D array and returns a single array.
- Ignore `ZeroDivisionError` errors, and leave `np.NaN`s as is.
- To verify your work, compare your outputted array to the actual `'depth'` column in `diamonds`.
- It may seem like `FunctionTransformer` is totally unnecessary, since we can compute depth percentages using broadcasting directly. However, as we will see in lecture, transformers can be **pipelined** with other processing steps which greatly simplifies our code.

```python

```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
diamonds = pd.read_csv(Path('data')/'diamonds.csv')
q6c_trans = TransformDiamonds(diamonds)
q6c_out = q6c_trans.transform_to_depth_pct(diamonds)
```

```python
grader.check("q6.3")
```

Output:
```
q6.3 results: All test cases passed!
```

## Congratulations! You're done with Lab 8! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To verify that all of your work is indeed in `lab.py`, and that you didn't accidentally implement a function in this notebook and not in `lab.py`, we've included another notebook in the lab folder, called `lab-validation.ipynb`. `lab-validation.ipynb` is a version of this notebook with only the `grader.check` cells and the code needed to set up the tests. 

### **Go to `lab-validation.ipynb`, and go to Kernel > Restart & Run All.** This will check if all `grader.check` test cases pass using just the code in `lab.py`.

Once you're able to pass all test cases in `lab-validation.ipynb`, including the call to `grader.check_all()` at the very bottom, then you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> best_transformation() in [1, 2, 3, 4]
True
```


Tests for q2 (points: None)
```
>>> set(out_q2.columns) == {'ordinal_cut', 'ordinal_clarity', 'ordinal_color'}
True
```
```
>>> np.unique(out_q2['ordinal_cut']).tolist() == [0, 1, 2, 3, 4]
True
```


Tests for q3 (points: None)
```
>>> out1_q3.shape == (53940, 20)
True
```
```
>>> out1_q3.columns.str.startswith('one_hot').all()
True
```
```
>>> out1_q3.isin([0,1]).all().all()
True
```
```
>>> out2_q3.shape[1] == 3
True
```
```
>>> out2_q3.columns.str.startswith('proportion_').all()
True
```
```
>>> ((out2_q3 >= 0) & (out2_q3 <= 1)).all().all()
True
```


Tests for q4 (points: None)
```
>>> out_q4.columns.str.contains(' * ').all()
True
```
```
>>> ('x * z' in out_q4.columns) or ('z * x' in out_q4.columns)
True
```
```
>>> out_q4.shape[1] == 15
True
```


Tests for q5 (points: None)
```
>>> len(out_q5) == 6
True
```
```
>>> isinstance(out_q5[0], numbers.Real)
True
```
```
>>> all(isinstance(x, str) for x in out_q5[2:-1])
True
```
```
>>> out_q5[1] > out_q5[-1]
True
```


Tests for q6.1 (points: None)
```
>>> isinstance(q6a_out, np.ndarray)
True
```
```
>>> q6a_out[172, 0] == 1
True
```
```
>>> q6a_out[0, 0] == 0
True
```


Tests for q6.2 (points: None)
```
>>> isinstance(q6b_out, np.ndarray)
True
```
```
>>> np.isclose(q6b_out[0, 0], 0.0075757, atol=0.0001)
True
```
```
>>> np.isclose(q6b_out[1, 0], 0.0025252, atol=0.0001)
True
```


Tests for q6.3 (points: None)
```
>>> len(q6c_out.shape) == 1
True
```
```
>>> np.isclose(q6c_out[0], 61.286, atol=0.0001)
True
```
