```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

# Lab 7 – Regular Expressions

## DSC 80, Fall 2023

### Due Date: Monday, November 20th at 11:59 PM

## Instructions
Welcome to the seventh lab assignment in DSC 80 this quarter!

Much like in DSC 10, this Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. Unlike DSC 10, the notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:
- **Unlike in DSC 10, labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the [course GitHub repository](https://github.com/dsc-courses/dsc80-2023-fa).
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- **To ensure that all of your work to be submitted is in `lab.py`, we've provided an additional uneditable notebook, called `lab-validation.ipynb`, that contains only the tests and their setup. Make sure you are able to run it top-to-bottom without error before submitting!**
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**Importing code from `lab.py`**:

* Below, we import the `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import os
from pathlib import Path
import re
```

## Question 1 – Practice with Regular Expressions 🛠
<div class="alert alert-block alert-warning">
<b>Note</b>: The functions in this question all have doctests in their docstrings. The doctests constitute the public tests; we will still run hidden tests on each of your functions. <b>Please don't change any of the docstrings!</b>
</div>

Regular expressions can be tricky, and the best way to gain familiarity with them is through lots of practice. In this question, you will work through ten exercises, each of which requires you to write a regular expression that matches strings that satisfy certain criteria. Make sure to take a close look at the doctests for each function in `lab.py`, as they provide useful guidance for the types of strings you should and shouldn't match.

***Notes:*** 
- Make sure to refer to the [Regular Expression Resources](https://dsc80.com/resources/#regular-expressions) posted on the course website. In particular, we recommend having [regex101.com](https://regex101.com/) open while working, along with the [cheat sheet](https://dsc80.com/resources/other/berkeley-regex-reference.pdf).

- Each exercise has a star rating, between 1 (⭐️) and 3 (⭐️⭐️⭐️) stars, indicating its difficulty level (1 being the easiest, 3 being the hardest). If you are spending lots of time on 1-star exercises, take a close look at the syntax from lecture, as there is probably an easier way of writing the necessary pattern!

<br>

### Exercise 1 (⭐️)

Write a regular expression that matches strings that have `'['` as the third character and `']'` as the sixth character.

<br>

### Exercise 2 (⭐️)

Write a regular expression that matches strings that are phone numbers that start with `'(858)'` and follow the format `'(xxx) xxx-xxxx'` (`'x'` represents a digit).

***Note:*** There is a space between `'(xxx)'` and `'xxx-xxxx'`.

<br>

### Exercise 3 (⭐️)

Write a regular expression that matches strings that:
- are between 6 and 10 characters long (inclusive),
- contain only alphanumeric characters, whitespace and `'?'`, and
- end with `'?'`.

<br>

### Exercise 4 (⭐️⭐️)

Write a regular expression that matches strings with exactly two `'$'`, one of which is at the start of the string, such that:
- the characters between the two `'$'` can be anything (including nothing) except the lowercase letters `'a'`, `'b'`, and `'c'`, (and `'$'`), and
- the characters after the second `'$'` can only be the **lowercase or uppercase** letters `'a'`/`'A'`, `'b'`/`'B'`, and `'c'`/`'C'`, with every `'a'`/`'A'` before every `'b'`/`'B'`, and every `'b'`/`'B'` before every `'c'`/`'C'`. There must be at least one `'a'` or `'A'`, at least one `'b'` or `'B'`, and at least one `'c'` or `'C'`.
    

<br>

### Exercise 5 (⭐️)
Write a regular expression that matches strings that represent valid Python file names, including the extension. 

***Note:*** For simplicity, assume that file names only contain letters, numbers, and underscores (`'_'`).

<br>

### Exercise 6 (⭐️)
Write a regular expression that matches strings that:
- are made up of only lowercase letters and exactly one underscore (`'_'`), and
- have at least one lowercase letter on both sides of the underscore.

<br>

### Exercise 7 (⭐️)
Write a regular expression that matches strings that start with and end with an underscore (`'_'`).

<br>

### Exercise 8 (⭐️)

Apple serial numbers are strings of length 1 or more that are made up of any characters, other than
- the uppercase letter `'O'`, 
- the lowercase letter `'i`', and 
- the number `'1'`.

Write a regular expression that matches strings that are valid Apple serial numbers.

<br>

### Exercise 9 (⭐️⭐️)

ID numbers are formatted as `'SC-NN-CCC-NNNN'`, where 
- SC represents state code in uppercase (e.g. `'CA'`),
- NN represents a number with 2 digits (e.g. `'98'`),
- CCC represents a three letter city code in uppercase (e.g. `'SAN'`), and
- NNNN represents a number with 4 digits (e.g. `'1024'`).

Write a regular expression that matches strings that are ID numbers corresponding to the cities of `'SAN'` or `'LAX'`, or the state of `'NY'`. Assume that there is only one city named `'SAN'` and only one city named `'LAX'`.

<br>

### Exercise 10 (⭐️⭐️⭐️)

Write a function named `match_10` that takes in a string and:
- converts the string to lowercase,
- removes all non-alphanumeric characters (i.e. removes everything that is not in the `\w` character class), and the letter `'a'`, and
- returns a list of every **non-overlapping** three-character substring in the remaining string, starting from the beginning of the string.
   
For instance, consider the following doctest:

```py
>>> match_10('Ab..DEF')
['bde']
```

Here's how `match_10` should process `'Ab..DEF'`:

1. Convert to lowercase: `'ab..def'`.
2. Remove non-alphanumeric characters and the letter `'a'`: `'bdef'`.
3. Starting from the beginning of the string, there is only a single non-overlapping three character substring: `'bde'`. Hence, we return `['bde']`.

***Note:*** Perform your operations in the exact order described above, otherwise your code may not pass all the tests. Don't use a `for`-loop.

```python
string = "ABCdef"
string = string.lower()
string = re.sub(r'[^\w]|a', '', string)
string[:3]
result = []
for i in range(0, len(string) - 2, 3):
    substring = string[i:i+3]
    if substring not in result:
      result.append(substring)
result
```

Output:
```
['bcd']
```

```python
match_1("ab[cd]ef")
```

Output:
```
True
```

```python
match_2('(858) 456-7890b')
```

Output:
```
False
```

```python
match_3('qwertsd?')
```

Output:
```
True
```

```python
match_6('aab_cbb_bc')
```

Output:
```
False
```

```python

```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Question 2 – Capturing Groups in Regular Expressions 📡

The dataset stored in `data/messy.txt` contains personal information from a fictional website that a user scraped from web server logs. Within this dataset, there are four fields that are of interest to you:
1. Email Addresses (assume they are alphanumeric usernames and domain names)
2. [Social Security Numbers](https://en.wikipedia.org/wiki/Social_Security_number#Structure)
3. Bitcoin Addresses (alphanumeric strings of long length)
4. Street Addresses

Complete the implementation of the function `extract_personal`, which takes in a string containing the contents of a server log file (like `open('data/messy.txt').read()`) and returns a **tuple of four separate lists** containing values of the 4 pieces of information listed above (in the order listed above). Do **not** keep empty values.

***Note:*** Since this data is messy, your function will be allowed to miss ~5% of the records in each list. Good spot checking using certain useful substrings (e.g. `'@'` for emails) should help assure correctness! Your function will be tested on a sample of the file `messy.txt`.

***Hint:*** There are multiple "delimiters" in use in the file; there are few enough of them that you can safely determine what they are.

```python
# experiment with extract_personal using the file s below
fp = Path('data') / 'messy.txt'
s = open(fp, encoding='utf8').read()
```

```python
re.findall(r'\b\w+@\w+\.\w+\b', s)
re.findall(r'\b\d{3}-\d{2}-\d{4}\b', s)
re.findall(r'\bbitcoin:\w+\b', s)
re.findall(r'\d{1,6} [A-Za-z0-9 ]+', s)
```

Output:
```
['814 Monterey Court',
 '62 Hooker Park',
 '27811 Clyde Gallagher Drive',
 '32553 Riverside Pass',
 '17157 Clemons Alley',
 '75 Dorton Parkway',
 '5 Manitowish Trail',
 '84 Northland Center',
 '64404 Sundown Street',
 '47144 Mockingbird Street',
 '413 Sutherland Court',
 '8906 Almo Lane',
 '34223 Graceland Crossing',
 '94 Mcguire Lane',
 '1 Delladonna Circle',
 '169 Claremont Point',
 '19475 Meadow Vale Crossing',
 '147 Cascade Center',
 '762 Knutson Terrace',
 '301 Monument Trail',
 '39 Buena Vista Alley',
 '94 Barby Circle',
 '98 Hansons Road',
 '83525 Calypso Way',
 '61 Heath Street',
 '60 Bay Hill',
 '923 Lindbergh Place',
 '1 Elmside Circle',
 '34 High Crossing Road',
 '63 Bonner Lane',
 '67 Superior Terrace',
 '846 Acker Road',
 '4 Menomonie Avenue',
 '858 Arrowood Terrace',
 '2566 Anhalt Trail',
 '30840 Sherman Hill',
 '29170 Parkside Lane',
 '150 Logan Street',
 '74 North Court',
 '51374 8th Hill',
 '6432 Continental Avenue',
 '46916 Donald Court',
 '1950 Beilfuss Hill',
 '949 
```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests
test_fp = Path('data') / 'messy.txt'
test_s = open(test_fp, encoding='utf8').read()
emails, ssn, bitcoin, addresses = extract_personal(test_s)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Question 3 – TF-IDF 📊

The dataset `data/reviews.txt` contains [Amazon reviews](https://cseweb.ucsd.edu/~jmcauley/datasets.html#amazon_reviews/) for ~200k phones and phone accessories. The dataset has already been "cleaned" for you. In this question, you will complete the implementation of a function, which takes in the reviews dataset as a Series (with one entry per review) as well as a single review, and returns the word that "best summarizes the single review" using TF-IDF.

To do so, implement the two functions below.

#### `tfidf_data`
You will first need to decompose the review into individual words, count their frequencies, and calculate the TF-IDF score for each word. Complete the implementation of the function `tfidf_data`, which takes in the reviews data as a Series (`reviews_ser`) and a single review (`review`) from that Series and returns a DataFrame indexed by the words in `review` with four columns:
- `'cnt'`: the number of times each word is found in the review (telling how frequently a word is used in the `review`)
- `'tf'`: the term frequency for each word (giving an idea of how important each word relative to the other words in the `review`)
- `'idf'`: the inverse document frequency for each word (helping understand how important each word is to the whole `reviews_ser`)
- `'tfidf'` the TF-IDF for each word (measuring how important each word is to `review`, taking into account both the frequency of the word in the `review` and across `reviews_ser`)

You may use a `for`-loop. The words in the outputted DataFrame may appear in any order.

***Hint:*** You may need to use the [`'\b'` character](https://www.regular-expressions.info/wordboundaries.html) somewhere.
    
<br>

#### `relevant_word`

Complete the implementation of the function `tfidf_data`, which takes in the DataFrame that `tfidf_data` returns and returns the word that "best summarizes" the review. If there are multiple "best" summary words, return any one of them.

```python
# experiment with tfidf_data using reviews_ser and review below 
fp = Path('data') / 'reviews.txt'
reviews_ser = pd.read_csv(fp, header=None).squeeze("columns")
review = open(Path('data') / 'reviews.txt', encoding='utf8').read().strip()
```

```python
reviews_ser
```

Output:
```
0        works great  i called t mobile and had this si...
1        these items looked to be of good quality and h...
2        this product arrive faster than i expected  i ...
3        i brought this for my sister who has a g2 but ...
4         i am both delighted and disappointed with del...
                               ...                        
58297    bought i for my husband's phone a week ago  ni...
58298    the spring loaded adjustable holder looks nift...
58299    received on time  everything working less spea...
58300    i bought this battery pack with the mindset of...
58301              that work    but thats not the original
Name: 0, Length: 58302, dtype: object
```

```python
len(reviews_ser[0].strip().split())
```

Output:
```
27
```

```python
cnt = pd.Series(reviews_ser[0].strip().split()).value_counts()
tf = cnt/ len(reviews_ser[0].strip().split())
idf = np.log(len(reviews_ser) / reviews_ser.str.contains('works').sum())
```

```python
reviews_ser[0].count(r'\bt\b')
cnt
```

Output:
```
a            2
works        1
great        1
to           1
speak        1
can          1
you          1
or           1
system       1
automated    1
have         1
they         1
minutes      1
5            1
in           1
activated    1
card         1
sim          1
this         1
had          1
and          1
mobile       1
t            1
called       1
i            1
rep          1
dtype: int64
```

```python
reviews_ser.str.split().str.len()
```

Output:
```
0          27
1          25
2          28
3          53
4        1787
         ... 
58297      22
58298     230
58299      35
58300     468
58301       7
Name: 0, Length: 58302, dtype: int64
```

```python
def calculate_tf(review):
    words = review.strip().split()
    word_series = pd.Series(words)
    tf = word_series.value_counts() / len(words)
    return tf
```

```python
tfidf_dict = {}
tf = calculate_tf(reviews_ser[0])
for word in tf.index:
    cnt = pd.Series(reviews_ser[0].strip().split()).value_counts()
    re_pat = fr'\b{word}\b'
    tf = len(re.findall(re_pat, reviews_ser[0].strip())) / len(reviews_ser[0].strip().split())
    idf = np.log(len(reviews_ser) / reviews_ser.str.contains(re_pat, regex=True).sum())
    tfidf_dict[word] = {'tf': tf, 'idf': idf, 'tfidf': tf * idf}

tfidf = pd.DataFrame.from_dict(tfidf_dict, orient='index')
tfidf['cnt'] = cnt
```

```python
tfidf
```

Output:
```
tf       idf     tfidf  cnt
a          0.074074  0.360392  0.026696    2
works      0.037037  1.898413  0.070312    1
great      0.037037  1.301647  0.048209    1
to         0.037037  0.382649  0.014172    1
speak      0.037037  5.505332  0.203901    1
can        0.037037  1.456449  0.053943    1
you        0.037037  1.096145  0.040598    1
or         0.037037  1.523326  0.056419    1
system     0.037037  4.633032  0.171594    1
automated  0.037037  8.083020  0.299371    1
have       0.037037  0.975412  0.036126    1
they       0.037037  1.719130  0.063671    1
minutes    0.037037  4.058661  0.150321    1
5          0.037037  2.271546  0.084131    1
in         0.037037  0.844683  0.031285    1
activated  0.037037  6.129205  0.227008    1
card       0.037037  3.866786  0.143214    1
sim        0.037037  4.989455  0.184795    1
this       0.037037  0.441295  0.016344    1
had        0.037037  1.689987  0.062592    1
and        0.037037  0.248188  0.009192    1
mobile     0.037037  4.0766
```

```python
# don't change this cell, but do run it -- it is needed for the tests
fp = os.path.join('data', 'reviews.txt')
reviews_ser = pd.read_csv(fp, header=None).squeeze("columns")
review = open(os.path.join('data', 'review.txt'), encoding='utf8').read().strip()
q3_tfidf = tfidf_data(reviews_ser, review)

try:
    q3_rel = relevant_word(q3_tfidf)
except:
    q3_rel = None
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Questions 4 and 5 – Tweet Analysis 🐥

The dataset `data/ira.csv` contains tweets tagged by Twitter as likely being posted by the [Internet Research Agency](https://en.wikipedia.org/wiki/Internet_Research_Agency), the tweet factory facing allegations for attempting to influence US political elections.

Questions 4 and 5 will focus on the following:
- Question 4: Look at the hashtags present in the text and trends in their makeup.
- Question 5: Prepare the dataset for modeling by creating features out of the text fields.

### Question 4 – Hashtags #️⃣

You may assume that a hashtag is any string without whitespace that immediately follows a `'#'`.

#### `hashtag_list`

Complete the implementation of the function `hashtag_list`, which takes in a Series of tweet texts and returns a Series containing a list of hashtags present in each tweet's text. If a tweet's text doesn't contain a hashtag, the Series should contain an empty list for that tweet. Don't include the `'#'` symbol in the lists that are returned.

<br>

#### `most_common_hashtag`

Complete the implementation of the function `most_common_hashtag`, which takes in a Series of hashtag lists (as is outputted by `hashtag_list`) and returns a Series consisting of a single hashtag per tweet: 
- If the tweet's text has no hashtags, the entry should in the output Series should be `NaN`.
- If the tweet's text has one distinct hashtag, the entry in the output Series should be that hashtag.
- If the tweet's text has more than one hashtag, the entry in the output Series should be the most common hashtag **in the tweet's text** with respect to **the whole input Series**. If there is a tie for the most common, any of the most common can be returned.
    - For example, if the input Series was `pd.Series([[2], [2], [3, 2, 3]])`, the output would be `pd.Series([2, 2, 2])`. Even though `3` was more common in the third list than `2`, `2` is more common than `3` among all hashtags in the Series.

```python
# The public tests don't test your work on the `ira` data,
# but the hidden tests do.
# So, make sure to thoroughly test your work yourself!
fp = os.path.join('data', 'ira.csv')
ira = pd.read_csv(fp, names=['id', 'name', 'date', 'text'])

```

```python
ira['text']
def extract_hashtags(tweet):
    return re.findall(r'#(\S+)', tweet)

re.findall(r'#(\S+)', ira['text'][4])


```

Output:
```
['IHatePokemonGoBecause', 'PokesAreJokes.']
```

```python
ira['text']
```

Output:
```
0        The Best Exercise To Lose Belly Fat In 2 weeks...
1        RT @Philanthropy: Dozens of ‘hate groups’ have...
2        Artificial intelligence can find, map poverty,...
3        Uber balks at rules proposed by world’s busies...
4        RT @dirtroaddiva1: #IHatePokemonGoBecause he  ...
                               ...                        
89995    Trump: Kasich shouldn't be allowed to run http...
89996    RT @JefLeeson: #ThingsYouCantIgnore The last s...
89997    RT @nealcarter: When someone said the first li...
89998    RT @indigenous01: #rantfortoday I speak the Wo...
89999    10 Things to Know for Monday https://t.co/XoOg...
Name: text, Length: 90000, dtype: object
```

```python
hashtag_list(ira['text'])
```

Output:
```
0        [Exercise, LoseBellyFat, CatTV, TeenWolf]
1                                               []
2                                           [tech]
3                                           [news]
4           [IHatePokemonGoBecause, PokesAreJokes]
                           ...                    
89995                                   [politics]
89996                        [ThingsYouCantIgnore]
89997                                           []
89998                               [rantfortoday]
89999                                       [news]
Name: text, Length: 90000, dtype: object
```

```python
text1 = hashtag_list(ira['text'])
most_common_hashtag(text1)
```

Output:
```
['Exercise', 'LoseBellyFat', 'CatTV', 'TeenWolf', 'tech', 'news', 'IHatePokemonGoBecause', 'PokesAreJokes', 'health', 'IWouldPreferToForget', 'NowPlaying', 'rap', 'hiphop', 'music', 'AthleticsTVShows', 'HillaryRottenClinton', 'TrumpTapes', 'entertainment', 'sports', 'politics', 'StopIslam', 'IslamKills', 'ParisAttacks', 'SlowJamTV', 'assaultatspringvalleyhigh', 'Brussels', 'WakeUpAmerica', 'OscarsSoWhite', 'oscarssoboring', 'OscarHasNoColor', 'SanJose', 'SextingWentWrongWhen', '2016In4Words', 'BlackHistoryMonth', 'MakeMeHateYouInOnePhrase', 'politics', 'AlternativeAcronymInterpretations', 'Castros', 'AgencySpy', 'racism', 'BTP', 'world', 'news', 'ObamasWishList', 'MustBeBanned', 'RunDakRun', 'DallasCowboys', 'business', 'politics', 'aid', 'disaster', 'environment', 'NowPlaying', 'HipHop', 'HillaryClinton', 'imwithher', 'StLouis', 'IfICouldntLie', 'news', 'morningjoe', 'news', 'MalcolmShabazz', 'LoveOfMyLife', 'AddAMovieRuinAMovie', 'NationalSelfieDay', 'baseball', 'WeAreTheMedia', 'pol
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

### Question 5 – Features 📋

Now, create a DataFrame of features from the `ira` data.  That is, create a function `create_features` that takes in a DataFrame `ira` that has just a single column, `'text'`, and returns a DataFrame with the same index as `ira` (i.e. the rows correspond to the same tweets) and the following columns:
* `'num_hashtags'`, the number of hashtags present in the tweet.
* `'mc_hashtags'`, the most common hashtag associated to the tweet (using the result of `most_common_hashtag` from Question 4).
* `'num_tags'`, the number of tags the tweet has (look for the presence of `'@'`).
* `'num_links'`, the number of hyperlinks present in the tweet.
    - A hyperlink is a string starting with `'http://'` or `'https://'`, not followed by whitespaces.
* `'is_retweet'`, a Boolean describing whether the tweet is a retweet. A retweet is a tweet that **begins** with `'RT'`.
* `'text'`, a version of the tweet's text that is cleaned according to the following steps, **in this exact order**:
    1. All meta-information above (retweet info, tags, hyperlinks, and hashtags) should be replaced with a single space.
    2. Everything other than letters, numbers, and spaces should be replaced with a single space.
    3. All letters should be lowercase.
    4. All words should be separated by exactly one space, and leading/trailing whitespace should be removed (stripped).
    
The columns in the outputted DataFrame must be in the order `['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']`. (Remember, the DataFrame that `create_features` is called on only has a single column, `'text'`.)

***Notes:***
- It's a good idea to make helper function for each column.
- The `\w` character class in regex **does not** refer to letters, numbers, and spaces (or even just letters and numbers). As such, you can't use it here!
- `create_features` will take a while to run on the entire dataset – test it on a small sample first!

```python
# The doctests/public tests don't test your work on the `ira` data,
# but the hidden tests do.
# So, make sure to thoroughly test your work yourself!
fp = os.path.join('data', 'ira.csv')
ira = pd.read_csv(fp, names=['id', 'name', 'date', 'text'])
```

```python
content = ira[['text']]
def num_hashtags(content):
    return len(re.findall(r'#[A-Za-z0-9]+', content))

# content['text'].apply(num_hashtags)
# ira['text'].apply(lambda x: most_common_hashtag(extract_hashtags(x))[0] if extract_hashtags(x) else None)
ira['text'].str.count('@[A-Za-z]+')
ira['text'].str.count('(http://\S+)|(https://\S+)')
ira['text'].str.startswith('RT') != 0
text = re.sub(r'@[A-Za-z]+|http[s]?://\S+|#[A-Za-z]+|RT', ' ', ira['text'][0])
text = re.sub(r'[^[A-Za-z0-9 ]+\s]', ' ', text)
text = text.lower()
text = re.sub(r'\s+', ' ', text)
text.strip()
```

Output:
```
'the best exercise to lose belly fat in 2 weeks …'
```

```python
# don't change this cell, but do run it -- it is needed for the tests
# (yes, we know it says "hidden" – there are still truly hidden tests in this question)
fp_hidden = 'data/ira_test.csv'
ira_hidden = pd.read_csv(fp_hidden, header=None)
text_hidden = ira_hidden.iloc[:, -1:]
text_hidden.columns = ['text']

test_hidden = create_features(text_hidden)
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

## Congratulations! You're done with Lab 7! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To verify that all of your work is indeed in `lab.py`, and that you didn't accidentally implement a function in this notebook and not in `lab.py`, we've included another notebook in the lab folder, called `lab-validation.ipynb`. `lab-validation.ipynb` is a version of this notebook with only the `grader.check` cells and the code needed to set up the tests. 

### **Go to `lab-validation.ipynb`, and go to Kernel > Restart & Run All.** This will check if all `grader.check` test cases pass using just the code in `lab.py`.

Once you're able to pass all test cases in `lab-validation.ipynb`, including the call to `grader.check_all()` at the very bottom, then you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!
```

Tests for q1 (points: None)
```
>>> import doctest
>>> doctest.run_docstring_examples(match_1, {'match_1': match_1})
>>> doctest.run_docstring_examples(match_2, {'match_2': match_2})
>>> doctest.run_docstring_examples(match_3, {'match_3': match_3})
>>> doctest.run_docstring_examples(match_4, {'match_4': match_4})
>>> doctest.run_docstring_examples(match_5, {'match_5': match_5})
>>> doctest.run_docstring_examples(match_6, {'match_6': match_6})
>>> doctest.run_docstring_examples(match_7, {'match_7': match_7})
>>> doctest.run_docstring_examples(match_8, {'match_8': match_8})
>>> doctest.run_docstring_examples(match_9, {'match_9': match_9})
>>> doctest.run_docstring_examples(match_10, {'match_10': match_10})
```
failure_message: doctests


Tests for q2 (points: None)
```
>>> emails[0] == 'test@test.com'
False
```
```
>>> ssn[0] == '423-01-9575'
False
```
```
>>> bitcoin[0] == '1BvBMSEYstWetqTFn5Au4m4GFg7xJaNVN2'
False
```
```
>>> addresses[0] == '530 High Street'
False
```


Tests for q3 (points: None)
```
>>> q3_tfidf['cnt'].sum() == 85
True
```
```
>>> 'before' in q3_tfidf.index
True
```
```
>>> q3_tfidf.loc['a', 'cnt'] == 4
True
```
```
>>> np.isclose(q3_tfidf.loc['phone', 'idf'], 0.820703, atol=0.5)
True
```
```
>>> q3_rel in q3_tfidf.index
True
```


Tests for q4 (points: None)
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = pd.DataFrame(testdata, columns=['text'])['text']
>>> out = hashtag_list(test)
>>> out.iloc[0] == ['NLP', 'NLP1', 'NLP1']
True
```
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = hashtag_list(pd.DataFrame(testdata, columns=['text'])['text'])
>>> most_common_hashtag(test).iloc[0] == 'NLP1'
True
```
```
>>> out = pd.Series([[2], [2], [3, 2, 3]])
>>> (most_common_hashtag(out) == 2).all()
True
```


Tests for q5 (points: None)
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = pd.DataFrame(testdata, columns=['text'])
>>> out = create_features(test)
>>> anscols = ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
>>> ansdata = [['text cleaning is cool', 3, 'NLP1', 1, 1, True]]
>>> ans = pd.DataFrame(ansdata, columns=anscols)
>>> (out == ans).all().all()
True
```
```
>>> test_hidden.columns.tolist() == ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
True
```
```
>>> 400 <= (test_hidden['num_tags'] == 0).sum() <= 450
True
```
```
>>> 450 <= (test_hidden['num_links'] == 0).sum() <= 550
True
```
